```python
from __future__ import print_function

%matplotlib inline
%precision 16

import numpy
import matplotlib.pyplot as plt
```

Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel $\rightarrow$ Restart) and then run all cells (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says YOUR CODE HERE or "YOUR ANSWER HERE", as well as your name and collaborators below:

# HW 1:  Review

## Question 1

**(a)** [5] Consider the polynomial
$$\begin{aligned}
    p(x) &= (x - 2)^9 \\
         &= x^9 - 18 x^8 + 144 x^7 - 672 x^6 + 2016 x^5 - 4032 x^4 + 5376 x^3- 4608 x^2+ 2304 x- 512.
\end{aligned}$$
Plot the two different representations of the same polynomial above using the points $x = 1.920, 1.921, 1.922, \ldots, 2.080$.

<!-- grade_id: cell-7314098343089560, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**(b)** [5] Compute each of the terms in the polynomial above separately (i.e. compute `x**9` by itself, `-18.0 * x**8`, etc.) for the value at $x = 2$ and some other $x$ values very close to $2$.

<!-- grade_id: cell-7411637601464065, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**(c)** [10] Describe the behavior you see in part (a) using what you saw in part (b).  If these are analytically equivalent what is the source of the differences?  Be as specific as you can be.

<!-- grade_id: cell-2221970078514962, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 2

**(a)** [15] Show how to derive analytically the backward finite difference approximation to the 1st derivative that is 2nd order accurate assuming a stencil that has equal width between the stencil points.  You can use the numerical method introduced in the lecture notes to check your work but you must show how to derive the coefficients analytically as well (you can use the procedure the function is based on).  Also make sure to explicitly find the leading order form of the truncation error $\tau$.

<!-- grade_id: cell-1951961903653690, points: 15, grade: true, solution: true -->
YOUR ANSWER HERE

**(b)** [10] Write a function that computes this 1st derivative of an arbitrary function $f(x)$ using the backward difference approximation you used above.  The function should accept a set of points `x` to evaluate the 1st derivative at and the function `f`.  For the first points in `x` use the forward difference approximation that mirrors the one you just found (see the lecture notes).

<!-- grade_id: cell-9997494412735662, points: 0, grade: false, solution: true -->
```python
def compute_prime(x, f):
    """Compute the first derivative"""
    # YOUR CODE HERE
    raise NotImplementedError()
    return f_prime
```

<!-- grade_id: cell-6147025779860527, points: 10, grade: true, solution: false -->
```python
f = lambda x: x**3 / numpy.sin(x)
f_prime = lambda x: -x**3 * numpy.cos(x) / numpy.sin(x)**2 + 3.0 * x**2 / numpy.sin(x)
x = numpy.linspace(0.25, 0.5 * numpy.pi, 150)
numpy.testing.assert_allclose(compute_prime(x, f), f_prime(x), atol=1e-3)
print("Success!")
```

**(c)** [15] Using your function above plot the derivative of the function and error convergence behavior for the functions.  
$$
    f(x) = \frac{x^3}{\sin x}, \quad \quad g(x) = |x - 3/2|, \quad \text{and} \quad h(x) = \left \{ \begin{aligned} 0 & & x < 3/2 \\ 1 & & x \ge 3/2 \end{aligned} \right .
$$
on the interval $x \in [1, 2]$.  Explain the behavior you see.

<!-- grade_id: cell-2489862782945804, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3579347744052423, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

# Question 3
Solving the following PDEs showing the relevant steps and plot a few time snapshots of the solution.

**(a)** [15] $u_t = \kappa u_{xx}$ with $x \in \Omega = (0, L)$, the boundary conditions $u(0, t)_x = 0$ and $u_x(L, t) = 0$, and the initial condition $u(x, 0) = 2 - \cos\left(\frac{\pi x}{L} \right) + 3 \cos\left(\frac{3 \pi x}{L} \right)$.  When plotting the solution pick a value for $L$ and play with values of $\kappa$.  Make sure your solution makes sense as $t \rightarrow \infty$.

<!-- grade_id: cell-2548330945171796, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-1316532658198281, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**(b)** [15] Solve the PDE $u_t + \left(\frac{1}{2}u^2\right)_x = 0$ where $x \in \Omega = \mathbb R$ and the initial conditions
$$ 
    u(x,0) = \left \{ \begin{aligned}
        1 & & x \leq 0 \\
        2 & & x > 0
    \end{aligned} \right .
$$
In addition to plotting the solution also mark the edges of interesting regions such as the rarefaction fan or the shock.

<!-- grade_id: cell-2036002999432267, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-1376908825388010, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```